# Week1 Assignment - Testing OKPy and the autograder.

For this assignment, you will load the iris flowers dataset from sklearn and predict flower species based on their petal and sepal sizes.

```python
"""This is a docstring!"""
```

$\beta=5$

alpha = 4

```python
%who
```

```python
var_list = %who_ls
```

```python
var_list
```

Output:
```
['datasets', 'iris']
```

## Part A. Load the Iris dataset.

Sklearn comes with prepared datasets that you can load. To load the iris dataset, run:
```python
from sklearn import datasets
iris = datasets.load_iris()
```

```python
from sklearn import datasets
iris = datasets.load_iris()
```

```python
import pandas as pd
```

```python
df = pd.DataFrame(iris.data, columns=iris.feature_names)
```

```python
df.head()
```

Output:
```
sepal length (cm)  sepal width (cm)  petal length (cm)  petal width (cm)
0                5.1               3.5                1.4               0.2
1                4.9               3.0                1.4               0.2
2                4.7               3.2                1.3               0.2
3                4.6               3.1                1.5               0.2
4                5.0               3.6                1.4               0.2
```

```python
df['sepal length (cm)'].head()
```

Output:
```
0    5.1
1    4.9
2    4.7
3    4.6
4    5.0
Name: sepal length (cm), dtype: float64
```

```python
df.columns = [s.replace(" ", "_") for s in df.columns]
```

```python
df.columns = [s.replace(")", "_").replace("(", "") for s in df.columns]
```

```python
df.columns
```

Output:
```
Index(['sepal_length_cm_', 'sepal_width_cm_', 'petal_length_cm_',
       'petal_width_cm_'],
      dtype='object')
```

```python
df.sepal_length_cm_.head()
```

Output:
```
0    5.1
1    4.9
2    4.7
3    4.6
4    5.0
Name: sepal_length_cm_, dtype: float64
```

```python
%lsmagic
```

Output:
```
Available line magics:
%alias  %alias_magic  %autoawait  %autocall  %automagic  %autosave  %bookmark  %cat  %cd  %clear  %colors  %conda  %config  %connect_info  %cp  %debug  %dhist  %dirs  %doctest_mode  %ed  %edit  %env  %gui  %hist  %history  %killbgscripts  %ldir  %less  %lf  %lk  %ll  %load  %load_ext  %loadpy  %logoff  %logon  %logstart  %logstate  %logstop  %ls  %lsmagic  %lx  %macro  %magic  %man  %matplotlib  %mkdir  %more  %mv  %notebook  %page  %pastebin  %pdb  %pdef  %pdoc  %pfile  %pinfo  %pinfo2  %pip  %popd  %pprint  %precision  %prun  %psearch  %psource  %pushd  %pwd  %pycat  %pylab  %qtconsole  %quickref  %recall  %rehashx  %reload_ext  %rep  %rerun  %reset  %reset_selective  %rm  %rmdir  %run  %save  %sc  %set_env  %store  %sx  %system  %tb  %time  %timeit  %unalias  %unload_ext  %who  %who_ls  %whos  %xdel  %xmode

Available cell magics:
%%!  %%HTML  %%SVG  %%bash  %%capture  %%debug  %%file  %%html  %%javascript  %%js  %%latex  %%markdown  %%perl  %%prun  %%pypy  %%
```

```python
%whos
```

Output:
```
Variable   Type         Data/Info
---------------------------------
datasets   module       <module 'sklearn.datasets<...>rn/datasets/__init__.py'>
df         DataFrame         sepal length (cm)  s<...>n\n[150 rows x 4 columns]
iris       Bunch        {'data': array([[5.1, 3.5<...>/datasets/data/iris.csv'}
pd         module       <module 'pandas' from '/U<...>ages/pandas/__init__.py'>
```

```python
%who_ls
```

Output:
```
['datasets', 'df', 'iris', 'pd']
```

```python
from pathlib import Path
```

```python
files = Path().glob("../*.py")
```

```python
list(files)
```

Output:
```
[PosixPath('../conf.py')]
```

```python
[f.read_text()]
```

```python
%%bash
echo hello world
```

Output:
```
hello world
```

```python
x = !echo hi
```

```python
x = !ls
```

```python
p = Path().home()
```

```python
p.absolute()
```

Output:
```
PosixPath('/Users/mohamedmm')
```

```python
f = p / "conf.py"
```

```python
f.read_text().splitlines([-1])
```

Output:
```
~/Documents/apps/anaconda3/lib/python3.7/pathlib.py in open(self, mode, buffering, encoding, errors, newline)
   1191             self._raise_closed()
   1192         return io.open(self, mode, buffering, encoding, errors, newline,
-> 1193                        opener=self._opener)
   1194 
   1195     def read_bytes(self):

~/Documents/apps/anaconda3/lib/python3.7/pathlib.py in _opener(self, name, flags, mode)
   1044     def _opener(self, name, flags, mode=0o666):
   1045         # A stub for the opener argument to built-in open()
-> 1046         return self._accessor.open(self, flags, mode)
   1047 
   1048     def _raw_open(self, flags, mode=0o777):

FileNotFoundError: [Errno 2] No such file or directory: '/Users/mohamedmm/conf.py'
```

```python
last = !tail -n 1 ../*.py
```

```python

```

```python
%%bash
find . -name "*.txt"
```

```python

```

```python
iris.target_names
```

Output:
```
array(['setosa', 'versicolor', 'virginica'], dtype='<U10')
```

```python
iris.target
```

Output:
```
array([0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0,
       0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0,
       0, 0, 0, 0, 0, 0, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1,
       1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1,
       1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2,
       2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2,
       2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2])
```

<!-- grade_id: cell-c1e7f79ad332b8f9, points: 0, grade: false, solution: true -->
```python
from sklearn import datasets

# Load the iris dataset
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-fd50146eeead0c65, points: 1, grade: true, solution: false -->
```python
assert (iris['data'] == datasets.load_iris()['data']).all()
```

## Part B. Use an SVM to predict iris species.

To setup the data and train an SVM, run the following:
```python
from sklearn.svm import SVC

X = iris.data
y = iris.target

classifier = SVC(kernel="rbf")
classifier.fit(X, y)
predictions = classifier.predict(X)
```

```python
%who
```

Output:
```
SVC	 X	 classifier	 datasets	 iris	 predictions	 var_list	 y
```

```python
df = pd.DataFrame(iris.data, columns=iris.feature_names)
```

```python
df %>%
    head
```

Output:
```
File "<ipython-input-56-ca597c5c3715>", line 1
    df %>%
        ^
SyntaxError: invalid syntax
```

```python
from sklearn.svm import SVC # support vector classifier
```

```python
# classification - outcome = discrete
# regression - outcome = continuous
```

```python
?SVC
```

Output:
```
[0;31mInit signature:[0m
[0mSVC[0m[0;34m([0m[0;34m[0m
[0;34m[0m    [0mC[0m[0;34m=[0m[0;36m1.0[0m[0;34m,[0m[0;34m[0m
[0;34m[0m    [0mkernel[0m[0;34m=[0m[0;34m'rbf'[0m[0;34m,[0m[0;34m[0m
[0;34m[0m    [0mdegree[0m[0;34m=[0m[0;36m3[0m[0;34m,[0m[0;34m[0m
[0;34m[0m    [0mgamma[0m[0;34m=[0m[0;34m'auto_deprecated'[0m[0;34m,[0m[0;34m[0m
[0;34m[0m    [0mcoef0[0m[0;34m=[0m[0;36m0.0[0m[0;34m,[0m[0;34m[0m
[0;34m[0m    [0mshrinking[0m[0;34m=[0m[0;32mTrue[0m[0;34m,[0m[0;34m[0m
[0;34m[0m    [0mprobability[0m[0;34m=[0m[0;32mFalse[0m[0;34m,[0m[0;34m[0m
[0;34m[0m    [0mtol[0m[0;34m=[0m[0;36m0.001[0m[0;34m,[0m[0;34m[0m
[0;34m[0m    [0mcache_size[0m[0;34m=[0m[0;36m200[0m[0;34m,[0m[0;34m[0m
[0;34m[0m    [0mclass_weight[0m[0;34m=[0m[0;32mNone[0m[0;34m,[0m[0;34m[0m
[0;34m[0m    [0mverbose[0m[0;34m=[0m[0;32mFalse[0m[0;34m,[0m[0;34m[0m
[0;34m[0m    [0mmax_iter[0m[0;34m
```

```python
SVC # shift tab
```

```python
# del :  remove variable
```

```python
classifier = SVC()
```

```python
X = iris.data
y = iris.target
# or y = pd.series(iris.target)
# shape is importNT RATHER THAN DATAFRAME OR ARRAY
```

```python

```

```python

```

```python
classifier = SVC()

X.shape
```

Output:
```
(150, 4)
```

```python

```

```python
classifier.fit(X, y)
```

Output:
```
/Users/mohamedmm/Documents/apps/anaconda3/lib/python3.7/site-packages/sklearn/svm/base.py:193: FutureWarning: The default value of gamma will change from 'auto' to 'scale' in version 0.22 to account better for unscaled features. Set gamma explicitly to 'auto' or 'scale' to avoid this warning.
  "avoid this warning.", FutureWarning)
SVC(C=1.0, cache_size=200, class_weight=None, coef0=0.0,
    decision_function_shape='ovr', degree=3, gamma='auto_deprecated',
    kernel='rbf', max_iter=-1, probability=False, random_state=None,
    shrinking=True, tol=0.001, verbose=False)
```

```python
predictions = classifier.predict(X)
```

```python
classifier.score(X,y)
```

Output:
```
0.9866666666666667
```

- Overfitting (Variance): tailored to a specific dataset (does not generalize well), misses the point, memorizes the data
- Underfitting (Bias): accuracy could be better, too general 

1. Import
2. Instantiate (a class)
  - Estimators:
    - Classifiers
    - Regressors
3. Fit
4. Predict

```python
predictions
```

Output:
```
array([0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0,
       0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0,
       0, 0, 0, 0, 0, 0, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1,
       1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 2, 1, 1, 1, 1, 1, 2, 1, 1, 1, 1,
       1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2,
       2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2,
       2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2])
```

```python
y == predictions
```

Output:
```
array([ True,  True,  True,  True,  True,  True,  True,  True,  True,
        True,  True,  True,  True,  True,  True,  True,  True,  True,
        True,  True,  True,  True,  True,  True,  True,  True,  True,
        True,  True,  True,  True,  True,  True,  True,  True,  True,
        True,  True,  True,  True,  True,  True,  True,  True,  True,
        True,  True,  True,  True,  True,  True,  True,  True,  True,
        True,  True,  True,  True,  True,  True,  True,  True,  True,
        True,  True,  True,  True,  True,  True,  True,  True,  True,
        True,  True,  True,  True,  True, False,  True,  True,  True,
        True,  True, False,  True,  True,  True,  True,  True,  True,
        True,  True,  True,  True,  True,  True,  True,  True,  True,
        True,  True,  True,  True,  True,  True,  True,  True,  True,
        True,  True,  True,  True,  True,  True,  True,  True,  True,
        True,  True,  True,  True,  True,  True,  True,  True,  True,
        True,  True,
```

```python

```

```python
from sklearn.model_selection import train_test_split
```

<!-- grade_id: cell-bc8ce57a60e5fd7a, points: 0, grade: false, solution: true -->
```python
from sklearn.svm import SVC

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-165c467d2a6e7dde, points: 1, grade: true, solution: false -->
```python
assert X.shape == (150, 4)
assert len(y) == 150
```